### CS 125 Assignment

Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Leo Wheeler"
COLLABORATORS = ""
```

---
---

# Lab 2

Objectives:

- Develop your ability to ask questions of the data using relational and logical operators.

---
## Created by:
- Michael Stobb
- With collaborators: 
    - None
- Lab2

Remember that you are **encouraged** to collaborate on lab activities – just be sure to: 
1. Document your collaborators and sources.
2. Don’t electronically share the code.
3. Understand what you submit. 


---

## Some Important Ideas

1. Earlier we were introduced to **Relational Operators**: `<` `<=` `>` `>=` `==` `!=`
    1. Work with both numeric and string data types
    2. Allow you to ask a question of the data and get a boolean (True/False) value as a response.
    3. The question that you are asking is about the relationship/order of the operands.


2. Here is another set of operators: **Logical Operators**: `and` `or` `not`
    1. Work with boolean values as operands
    2. Form more complex questions (usually with different variables) and return a boolean value.

## (0) Small Experiments
Remember that these small experiments are designed to help you get practice with some of the concepts associated with the lab. They also should help you understand some of the nuances and peculiarities that you may encounter
when programming. Make sure that you work through **all** of the examples. If something unexpected happens, try
to make sense of it.  Above all, be *thinking*.  If you can’t make sense of it, ask a friend. If neither of you knows what is happening, talk Instructor (via email, Google Hangouts, or Office Hours).

1. Recall that floating point values sometimes lack precision. That is to say the way that the machine stores these values can introduce very small but very real rounding errors. Try this:
    1. ```python
    num = 1/2```
    2. ```python
    print ( num * 2)```
    3. ```python
    num = 1/49```
    4. ```python
    print (num * 49)```

```python
num = 1/2
```

```python
print(num*2)
```

Output:
```
1.0
```

```python
num = 1/49
```

```python
print(num*49)
```

Output:
```
0.9999999999999999
```

That isn’t very satisfying, but the human brain can understand that .999999999999999 is really, *really* close
to 1. But what if I try to do a comparison using a relational operator? 
- Write an expression using a relational operator that asks if num multiplied by 49 is the same as 1.

```python
abs(num*49 -1)<0.00001
```

Output:
```
True
```

It turns out that using relational operators with floats (especially ==) can be problematic. Try this
alternative:
- ```python
abs(num *49 – 1) <.00001```


This is a programmatic way to effectively say that the difference between these two values is “close enough”.
- `abs()` is a function that computes the absolute value.
- Can you explain where “close enough” comes from?

### - The function is showing that the difference between `num*49` and 1 is less than 0.00001

2. Relational operators work nicely with text – try the following
    1. ```python
    "a" < "b"```
    2. ```python
    "a" < "A"```
    3. ```python
    "cat" < "dog"```
    4. ```python
    "cat" < "Dog"```
    5. ```python
    "cat" < "cab"```

```python
"a"<"b"
```

Output:
```
True
```

```python
"a"<"A"
```

Output:
```
False
```

```python
"cat"<"dog"
```

Output:
```
True
```

```python
"cat" < "Dog"
```

Output:
```
False
```

```python
"cat"<"cab"
```

Output:
```
False
```

3. Remember that sometimes numeric data is stored as text – watch what happens and make sure you understand:
   1. ```python
    5 < 6```
   2. ```python
    55 < 6```
   3. ```python
    "5" <"6"```
   4. ```python
    "55" < "6"```

Don’t forget to cast the “numeric” strings if you want to use them as numbers.

```python
5<6
```

Output:
```
True
```

```python
55<6
```

Output:
```
False
```

```python
"5"<"6"
```

Output:
```
True
```

```python
"55"<"6"
```

Output:
```
True
```

4. Be careful when you use `==` with string operators:
    1. ```python
    capital = "Des Moines"```
    2. ```python
    capital == "DesMoines"```
    3. ```python
    capital == "Des moines"```
    4. ```python
    capital == "Des Moines"```

```python
capital = "Des Moines"
```

```python
capital == "DesMoines"
```

Output:
```
False
```

```python
capital == "Des moines"
```

Output:
```
False
```

```python
capital == "Des Moines"
```

Output:
```
True
```

5. Logical operators allow you to construct more complex expressions by combining Boolean values. Consider the following expressions:
    1. ```python
    p = True```
    2. ```python
    q = False```
    3. ```python
    p and q```
    4. ```python
    p or q```
    5. ```python
    not p```

```python
p = True
```

```python
q = False
```

```python
p and q
```

Output:
```
False
```

```python
p or q
```

Output:
```
True
```

```python
not p
```

Output:
```
False
```

Since there are only two Boolean values (`True` and `False`) we can exhaustively list out the outcomes of the logical operators. Complete the following tables by altering the values of `p` and `q` as operands to the logical operators.  Complete them IN THIS CELL.  Just double click here and you can edit directly.  Once you're finished, you can Shift+Enter to reformat the cell nicely.

1. 


| `p`   | `q`   | `p and q` |
|-------|-------|-----------|
| True  | True  |   True    |
| True  | False |  False    |
| False | True  |   False   |
| False | False | False     |


2. 


| `p`   | `q`   | `p or q` |
|-------|-------|-----------|
| True  | True  |  True     |
| True  | False |  True     |
| False | True  |  True     |
| False | False |  False    |



3. 


| `p`   | `not p`|
|-------|--------|
| True  | False     |
| False | True   |

---
# Lab Assignment

For each of the following activities, you are responsible for completing the activity, sharing your work with the 
Instructor or TA and submitting a final version to Moodle.

### Problem 0

The following function asks a question and returns an answer.  Specifically, it asks, *Did I work overtime?*.  Overtime pay is mandatory when a worker exceeds 40 hours of work per week.

```python
def madeOvertime(hours):
    return hours > 40
```

After defining the function, I want to test it, **multiple times**, to make sure it is working as expected.

```python
madeOvertime(30)
```

Output:
```
False
```

```python
madeOvertime(41)
```

Output:
```
True
```

```python
madeOvertime(40)
```

Output:
```
False
```

```python
madeOvertime(50)
```

Output:
```
True
```

As another, slightly more sophiticated example, I can ask the question, *Will Amazon make a new season of "The Boys"?* (Hint: They just released it!).  This depends on the following:
    - The ratings need to be above 80
    - The budget needs to be below 10,000,000
    
I could then write the following code:

```python
def nextSeason(ratings,budget):
    return (ratings > 80) and (budget < 10000000)
```

```python
nextSeason(75,8000000)
```

Output:
```
False
```

```python
nextSeason(85,8000000)
```

Output:
```
True
```

```python
nextSeason(75,11000000)
```

Output:
```
False
```

```python
nextSeason(85,11000000)
```

Output:
```
False
```

Make sure you fully understand the above examples before moving on to the next section.

Your job is to **edit the following code** and to write a function that provides the answer to the desired question. Your code will allow other programmers to use your function to supply input parameters and get an answer to the desired question. Do not ask the user for input directly, supply it as function inputs.  In each case, you will return a True/False answer to their question, just like the examples above.


For those of you who have done a little programming: you may be tempted to use an "if" statement. **DO NOT!** This activity is focused on asking the question, not processing the answer. Every answer *can and should* be the result of a simple boolean expression.


I would suggest that you write one function at a time – then convince yourself that it works by running
some test cases as demonstrated above.

---

1. Did I make a profit?
    - Profit is any amount remaining after you have paid your expenses.

```python
def makeProfit(earnings, expenses):
    
    return earnings > expenses
```

```python
makeProfit(500,340)
```

Output:
```
True
```

```python
makeProfit(500,650)
```

Output:
```
False
```

<!-- grade_id: cell-06557dbcf9c84b8a, points: 1, grade: true, solution: false -->
```python
assert makeProfit(100, 20) == True
assert makeProfit(50, 200) == False
```

2. Does the patient have a fever?
    - The normal body temperature is 98.6, but some variations exist.
    - The patient has a fever if the body temperature is greater than 100 degrees.

```python
def hasFever(temp):
    
    return temp >= 100
```

```python
hasFever(99)
```

Output:
```
False
```

```python
hasFever(102)
```

Output:
```
True
```

<!-- grade_id: cell-9af19748eb90d3a9, points: 1, grade: true, solution: false -->
```python
assert hasFever(20) == False
assert hasFever(200) == True
```

3. Did I get the wrong answer?
    - I am took a quiz and I want to know how I did.
    - Note: My answer and the correct answer should be STRINGS, like "A", "B", "C", etc.

```python
def answerWrong(my_ans, correct_ans):
    "A" == "A"
    "B" == "B"
    "C" == "C"
    "D" == "D"
    return my_ans == correct_ans
```

```python
answerWrong("A", "B")
```

Output:
```
False
```

```python
answerWrong("C","C")
```

Output:
```
True
```

<!-- grade_id: cell-18749c38b4c383f3, points: 1, grade: true, solution: false -->
```python
assert answerWrong("A","A") == True
assert answerWrong("C","D") == False
```

---
4. Did I pass?
    - 60% and above is considered a "passing" grade.
    - Grade should be entered as a number.

```python
def pass_test(score):
    
    
    return score >= 60
```

```python
pass_test(85)
```

Output:
```
True
```

```python
pass_test(45)
```

Output:
```
False
```

```python
pass_test(60)
```

Output:
```
True
```

<!-- grade_id: cell-c3a0029ec31d87ce, points: 1, grade: true, solution: false -->
```python
assert pass_test(90) == True
assert pass_test(30) == False
```

5. Is my letter a capital letter?
    - Input should be a single letter (as a string!)

```python
def is_capital_letter(my_letter):
    
    return my_letter < "a" and my_letter >= "A"
```

```python
is_capital_letter("f")
```

Output:
```
False
```

```python
is_capital_letter("A")
```

Output:
```
True
```

<!-- grade_id: cell-8c0bf7a844a23a50, points: 1, grade: true, solution: false -->
```python
assert is_capital_letter("A") == True
assert is_capital_letter("f") == False
```

6. Is this letter a vowel?
    - Input should be a single letter (as a string)
    - Hint, use lots of `and`s

```python
def is_vowel(my_letter):
      
    return (my_letter == "a") or (my_letter == "e") or (my_letter == "i") or (my_letter == "o") or (my_letter == "u") or (my_letter == "A") or (my_letter == "E") or (my_letter == "I") or (my_letter == "O") or (my_letter == "U")
```

```python
is_vowel("u")
```

Output:
```
True
```

```python
is_vowel("g")
```

Output:
```
False
```

<!-- grade_id: cell-e5db0b1ae9721aa1, points: 1, grade: true, solution: false -->
```python
assert is_vowel("A") == True
assert is_vowel("f") == False
assert is_vowel("e") == True
```

7. Is this number even?
    - Even numbers are evenly disible by 2.

```python
def is_number_even(N):
    
    return (N % 2) == 0
```

```python
is_number_even(4)
```

Output:
```
True
```

```python
is_number_even(7)
```

Output:
```
False
```

<!-- grade_id: cell-fa9cd2dee38bc07a, points: 1, grade: true, solution: false -->
```python
assert is_number_even(2) == True
assert is_number_even(7) == False
assert is_number_even(2022) == True
```

8. is my GMR abnormal?
    - GMR is some important measurement.
    - Normal values for GMR are between 15 and 48.

```python
def normal_GMR(my_gmr_value):
    
    return my_gmr_value > 15 and my_gmr_value <48
```

```python
normal_GMR(43)
```

Output:
```
True
```

```python
normal_GMR(12)
```

Output:
```
False
```

<!-- grade_id: cell-6040f0397aec8a44, points: 1, grade: true, solution: false -->
```python
assert normal_GMR(20) == True
assert normal_GMR(50) == False
assert normal_GMR(35) == True
```

9. Is this student a Sophomore?
    - Students are classified according to the number of credits that they have completed.
        - A Freshman is any student with less than 8 credits.
        - With every 8 credits students are successively ranked as Sophomores and Juniors.
        - Students with 24 or more credits are Seniors.

```python
def is_sophomore(num_credits):
    
    
    return num_credits >= 8 and num_credits <16
```

```python
is_sophomore(12)
```

Output:
```
True
```

```python

```

<!-- grade_id: cell-4560a4622fb49aba, points: 1, grade: true, solution: false -->
```python
assert is_sophomore(5) == False
assert is_sophomore(10) == True
assert is_sophomore(23) == False
assert is_sophomore(25) == False
```

10. Is it a good day to golf?
    - The best days for golf are influenced by the weather.
        - Golfing in the rain is no fun.
        - It is better to golf when it isn't too windy (windspeed less than 20 mph).
        - The ideal temperature is between 60 and 85 degrees.

```python
def good_golf_day(rain_status, windspeed, temperature):
    
    
    return (rain_status == False and windspeed < 20 and temperature <85 and temperature >60)
```

```python
good_golf_day(False, 10, 75)
```

Output:
```
True
```

```python
good_golf_day(False, 40,100)
```

Output:
```
False
```

<!-- grade_id: cell-1be57321bc244409, points: 1, grade: true, solution: false -->
```python
assert good_golf_day(False, 10, 75) == True
assert good_golf_day(True, 10, 75) == False
assert good_golf_day(False, 25, 75) == False
assert good_golf_day(False, 10, 90) == False
```

11. Can I retire?
    - IPERS is the Iowa Public Employees Retirement System.  To be eligible to collect pay from IPERS (i.e. to retire), you must meet one of three requirements.
        1. the sum of the age and years of service equals or exceeds 88
        2. age is at least 62 and years service is at least 20
        3. age is 65 or more

```python
def can_retire(current_age, years_of_service):

    return (current_age+years_of_service >= 88) or (current_age >=62 and years_of_service >= 20) or (current_age >=65)
```

```python
can_retire(62,19)
```

Output:
```
False
```

```python
can_retire(62,20)
```

Output:
```
True
```

<!-- grade_id: cell-ed500cbf60c75c60, points: 1, grade: true, solution: false -->
```python
assert can_retire(55, 10) == False
assert can_retire(50, 40) == True
assert can_retire(64, 20) == True
assert can_retire(70, 10) == True
```

12. Is my blood sugar level ok?
    - Normal Blood sugar levels are 70-100 when fasting, but 70-140 otherwise.
    - Fasting means that it has been more than 24 hours since you have eaten.

```python
def normal_blood_sugar(is_fasted, blood_sugar):
    
    
    return ((is_fasted == False) and blood_sugar >=70 and blood_sugar <=140) or ((is_fasted == True) and blood_sugar>=70 and blood_sugar<=100)
```

```python
normal_blood_sugar(True, 90)
```

Output:
```
True
```

```python
normal_blood_sugar(False, 150)
```

Output:
```
False
```

<!-- grade_id: cell-add51f8dbb2b559c, points: 1, grade: true, solution: false -->
```python
assert normal_blood_sugar(False, 130) == True
assert normal_blood_sugar(True, 80) == True
assert normal_blood_sugar(True, 110) == False
assert normal_blood_sugar(False, 150) == False
```

### Problem 1
We typically think of a year consisting of 365 days. In fact, it takes the Earth approximately 365.252375 days to go around the sun. To compensate for the fractional days, we have introduced the concept of a leap year. We typically think of this occurring every 4 years. But if we do that, you will notice that we are still undercompensating. (what happens to the extra .002375?!) 

It turns out the precise way to know if a year is a Leap year is as follows: 
- Leap years are evenly divisible by 4
- unless it can be evenly divisible by 100 
- except if it can be evenly divisible by 400.

So consider the following years – were they leap years?

- 2020, 2016, 1972, 2100
- 2012, 1974, 1832, 1644
- 2000, 1900, 1950, 1240

Write a function called `is_leap_year(year)`. This function should return `True` when the year provided is a leap year and `False` otherwise.

```python
def is_leap_year(year):
    year = int(year)
    
    return ((year/4 == int(year/4)) and (((year/100 == int(year/100)) == False ) or ((year/100 == int(year/100)) == True )and (year/400 == int(year/100))))
```

```python
is_leap_year(2100)
```

Output:
```
False
```

```python
is_leap_year(2020)
```

Output:
```
True
```

<!-- grade_id: cell-0191779c8bb25fb5, points: 1, grade: true, solution: false -->
```python
assert is_leap_year(1804) == True
assert is_leap_year(2020) == True
assert is_leap_year(2100) == False
```

```python

```