### CS 125 Assignment

Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Leo Wheeler"
COLLABORATORS = ""
```

---
---

# Lab 3

Objectives:

- Develop code that executes conditionally
- Recognize Selection patterns:
    - Alternate Value/Action
    - Cascading If/Classification
- Experiment with functions

---
## Created by:
- Michael Stobb
- With collaborators: 
    - None

Remember that you are **encouraged** to collaborate on lab activities – just be sure to: 
1. Document your collaborators and sources.
2. Don’t electronically share the code.
3. Understand what you submit. 


---

## Some Important Ideas

An if statement allows code to be executed - or not - based on the result of a boolean expression. When the condition evaluates to True, then the statements of the indented body are executed. When the condition is False, then the body is skipped. An `else` clause provides an alternative action that can be executed when the condition is False, as demonstrated below (**you don't need to run these**):

```python
if COND:
    command1
    command2
else:
    command1
    command2
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-de653bbc03c5> in <module>
----> 1 if COND:
      2     command1
      3     command2
      4 else:
      5     command1

NameError: name 'COND' is not defined
```

Here's two examples conditional statements (**you don't need to run these**):

```python
# Example 1
if (num != 0):
    avg = (amt/num)
    
    
# Example 2
if (num % 2 == 0):
    num_type = "Even"
else:
    num_type = "Odd"
```

## Example Program

Below are some functions to calculate grades.  The test function compares the expected output from the grade function to the actual output. Unfortunately, the grade function contains a logic error (similar to one that we saw in class), so our tests all report `False`. 

You should:
1. Fix the logic error in the grade function. (This should make your actual and expected outputs match)
2. Add at least three additional test cases to the test function.

```python
def grade(score):
    if score >= 90:
        letterGrade = "A"
    elif 80 <= score:
        letterGrade = "B"
    elif 70 <= score:
        letterGrade = "C"
    elif 60 <= score:
        letterGrade = "D"
    else:
        letterGrade = "F"
    return letterGrade
```

```python
assert grade(90) == "A"
assert grade(84) == "B"
assert grade(72) == "C"
assert grade(43) == "F"
```

```python

```

---
# Lab Assignment

### Problem 0

Acme Inc. assigns shipping charges at \\$1.25 per pound, or a flat rate of \\$2 for all orders less than 3 lbs. Prompt the user to enter the weight of their shipment using a main function – then pass that value to a function which computes the appropriate shipping charges. **You should make 2 functions to answer this question:** one that asks for the weight and prints the answer, and a seperate function which computes the charges.

Use your program to answer the following test cases:


| Input           | Output |
|-----------------|--------|
| Weight = 1lb    |        |
| Weight = 1.5lbs |        |
| Weight = 3lbs   |        |
| Weight = 10lbs  |        |

```python
def compute_shipping_cost(weight):
    if weight < 3:
        price = 2.00
    else:
        price = (1.25*weight)
    return price
```

```python
compute_shipping_cost(6.3)
```

Output:
```
7.875
```

```python
compute_shipping_cost(4)
```

Output:
```
5.0
```

<!-- grade_id: cell-770faf4103a799ea, points: 1, grade: true, solution: false -->
```python
assert compute_shipping_cost(1) == 2.00
assert compute_shipping_cost(2.5) == 2.00
assert compute_shipping_cost(5) == 6.25
```

```python
def determine_customer_cost():
    consumerPrice = compute_shipping_cost(int(input("Enter the weight of the package:")))
    print("Cost of shipping is: $", consumerPrice)
    return
```

```python
determine_customer_cost()
```

Output:
```
Enter the weight of the package:7
Cost of shipping is: $ 8.75
```

```python
determine_customer_cost()
```

Output:
```
Enter the weight of the package:185
Cost of shipping is: $ 231.25
```

### Problem 1
A certain profesor's class has four exams that each count for 25% of your final grade. However, if a student can show a 10% improvement on the final (fourth) exam over the average of their first three exams, the professor will count the final exam score as 100% of the final grade. For example, if a student averaged 75% on the first three exams, but earned 83% on the final exam; the final grade would be an 83%. Calculate a student‟s grade based on their exam scores.  **This should be answered with a single function**

Consider the following sample test cases. It probably makes sense to compute these by hand and then test them with your function.

| Input                                      | Output |
|--------------------------------------------|--------|
| Exam1: 75; Exam2: 75; Exam3: 75; Exam4: 75 |        |
| Exam1: 75; Exam2: 75; Exam3: 75; Exam4: 83 |        |
| Exam1: 75; Exam2: 75; Exam3: 75; Exam4: 82 |        |
| Exam1: 75; Exam2: 85; Exam3: 80; Exam4: 88 |        |
| Exam1: 75; Exam2: 75; Exam3: 85; Exam4: 85 |        |

```python
def calculate_final_grade(exam1, exam2, exam3, exam4):
    if ((exam1+exam2+exam3)/3 <= (exam4 - 1/10)):
        final_score = exam4
    else:
        final_score = int((exam1+exam2+exam3+exam4)/4)
    
    return final_score
```

```python
calculate_final_grade(60,70,78, 94)
```

```python

```

```python

```

```python

```

<!-- grade_id: cell-a88f161667b4ed33, points: 1, grade: true, solution: false -->
```python
assert calculate_final_grade(75,75,75,75) == 75.0
assert calculate_final_grade(75,75,75,83) == 83.0
```

### Problem 2
The wavelength of visible light ranges from 380 to 750 nanometers(nm).  While the spectrum is continuous, we often think of light as being characterized by 6 distinct colors :

| Color | Wavelength |
|-------|------------|
|violet |380-450 nm|
|blue |450-495 nm|
|green |495-570 nm|
|yellow |570-590 nm|
|orange |590-620 nm|
|red |620-750 nm|

**Write a single function** that accepts the wavelength and returns the appropriate color (as a string). If an "invalid" wavelength is provided, the function should return: `outside visible spectrum`.  Let the edge cases (450nm, 495nm, etc.), be part of the lower color, so 450nm is violet, and 495nm is blue, etc.

```python
def determine_color(wavelength):
    w = wavelength
    if w < 380:
        w = "outside visible spectrum"
    elif w >= 380 and w <= 450:
        w = "violet"
    elif w >450 and w<=495:
        w = "blue"
    elif w > 495 and w <= 570:
        w = "green"
    elif w > 570 and w <= 590:
        w = "yellow"
    elif w > 590 and w <= 620:
        w = "orange"
    elif w > 620 and w <= 750:
        w = "red"
    elif w > 750:
        w = "outside visible spectrum"
    
    return w
```

```python
determine_color(390)
```

Output:
```
'violet'
```

```python
determine_color(3)
```

Output:
```
'outside visible spectrum'
```

```python
determine_color(630)
```

Output:
```
'red'
```

<!-- grade_id: cell-eef112086379de4a, points: 1, grade: true, solution: false -->
```python
assert determine_color(390) == "violet"
assert determine_color(595) == "orange"
assert determine_color(570) == "green"
assert determine_color(800) == "outside visible spectrum"
```

### Problem 3
Given the lengths of three legs of a potential triangle, write a function to classify the type of triangle that results:

- Equilateral (all sides are the same length),
- Isosceles (exactly two sides have the same length),
- Scalene (all sides are different lengths) or
- Invalid (the sum of the two smaller sides are not longer than the longest side. See hint below 1 )


Note: your program should accept the legs of the triangle **in any order**. For example:

```python
triangle_classify(3,4,4)```
```python
Isoceles```



```python
triangle_classify(4,3,4)```
```python
Isoceles```

Be sure to test function and demonstrate that your code works (i.e. the expected output matches the actual output for a variety of cases.

**BONUS:** After you get the program working, try to add some code to also report when the triangle is acute, right, or obtuse.

```python
def triangle_classify(side1, side2, side3):
    x = int(side1)
    y = int(side2)
    z = int(side3)
    if (x == y) and (y == z):
        tri_type = "Equilateral"
    elif ((x == y) and z < y) or ((x == z) and y < z) or ((y == z) and x < y):
        tri_type = "Isoceles"
    elif ((x + y) < z) or ((y + z) < x) or ((x + z) < y):
        tri_type = "Invalid"
    else:
        tri_type = "Scalene"
    
    return tri_type
```

```python
triangle_classify(6,6,8)
```

Output:
```
'Scalene'
```

```python
triangle_classify(12,12,12)
```

Output:
```
'Equilateral'
```

```python
triangle_classify(13,13,135)
```

Output:
```
'Invalid'
```

<!-- grade_id: cell-ce4c8cfd85bbcf07, points: 1, grade: true, solution: false -->
```python
assert triangle_classify(3,4,4) == "Isoceles"
assert triangle_classify(4,3,4) == "Isoceles"
assert triangle_classify(4,4,3) == "Isoceles"
assert triangle_classify(3,3,3) == "Equilateral"
assert triangle_classify(3,4,5) == "Scalene"
assert triangle_classify(2,2,6) == "Invalid"
```

```python

```